# This is your first assignment

<!-- grade_id: cell-438c4f4a0e5ea709, points: 1, grade: true, solution: true -->
```python
""" Write your code below """

def helloWorld():
    """Prints 'hello world'"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e478356121f47838, points: 2, grade: true, solution: false -->
```python
""" Here is a sample of the tests your function should pass """
assert(helloWorld() == None)
```